# Week 4 Exercises

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 15 through 20 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 20.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names separated by commas.

```
sort_names(['Paul', 'Boal'],['Kermit','the','Frog'],['Donald','Duck'])

'Paul Boal, Donald Duck, Kermit the Frog'
```

```python
def sort_names(names):
    
    sorted_names = ""
    ### BEGIN SOLUTION
   
    formatted_names=[]
    for name in names:
        if len(name)==3:
            sortable_name=name[2]+ ', '+ name [1] + ' '+name[0]
            printable_name=name[0]+ ', '+ name [1] + ' '+name[1]
    else:
        sortable_name= name[1] + ','+ name [0]
        printable_name= name[0] + ','+ name [1]
    formatted_names.append([sortable_name, printable_name])
    formatted_names.sort()         
    
    for name in formatted_names:
        if sorted_names=='':
            sorted_name= name[1]
        else:
            sorted_names += ', ' + name[1]
 

    ### END SOLUTION
    
    return sorted_names
```

```python
names=  ['Paul', 'Boal'],['Kermit','the','Frog'],['Donald','Duck']
formatted_names=[]
for name in names:
    if len(name)==3:
        sortable_name=name[2]+ ', '+ name [1] + ' '+ name[0]
        printable_name=name[0]+ ', '+ name [1] + ' '+name[1]
    else:
        sortable_name= name[1] + ','+ name [0]
        printable_name= name[0] + ' '+ name [1]
    formatted_names.append([sortable_name, printable_name])
formatted_names.sort()         
output=''    
for name in formatted_names:
    output=output+', ' + name[1]
print(output)
```

```python
names = ['Paul', 'Boal'],['Kermit','the','Frog'],['Donald','Duck']

# 1. Create a list that I can easily sort
#    But keep a key to the original input so I can go back to that later
sortable_names = []
for key, name in enumerate(names):
    if len(name) == 3:
        sortable_name = '{}, {} {}'.format(name[2],name[0],name[1])
    else:
        sortable_name = '{}, {}'.format(name[0],name[1])
    sortable_names.append((sortable_name, key))

# 2. Sort that list
sorted_names = sorted(sortable_names)
sorted_names
```

```python
# 3. Format that sorted names the way I want them
pretty_names = []
for name in sorted_names:
    name_key = name[1]
    name_parts = names[name_key]
    pretty_name = ' '.join(name_parts)
    pretty_names.append(pretty_name)

pretty_names
```

```python
# 4. Combine names into an output string
output = ', '.join(pretty_names)
output
```

### 20.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Elliie'],
 ['Lester-Boal', 'Ada'],
 ...
```

```python
Families= {
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}
```

```python
output= []
for last_name, first_names in Families.items():
    print(last_name + '--'+ str(first_names))
    for first in first_names:
        output.append([last_name, first])
print(output)
    
```

```python
def dict_to_list(families):
    
    output = []
            
    return output
```

```python
def dict_to_list(families):
    
    output = []
    for family, names in families.items():
        for name in names:
            output.append([family, name])
            
    return output
```

```python
dict_to_list({
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
})
```

```python

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.
```

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():

    names = [['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    # To help you, I've provided a blank dictionary named `families`.
    families = {}
    # Your code should loop through `names` and do something to build
    # up `families` to be the output dictionary shown in the test
    # cells below.

    ### BEGIN SOLUTION
    for names in names:
        if names[0] in families:
            families[names[0]].append(names[1])
        else:
            families[names[0]] = [names[1]]
    ### END SOLUTION

    return families

```

```python
list_to_dict()
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict() == answer)
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

```python

```

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):
    conclusions = []
    ### BEGIN SOLUTION
    import pandas as pd
    avg_los = {
                "Hemolytic jaundice and perinatal jaundice" : 2,
                "Medical examination/evaluation" : 3.2,
                "Liveborn" : 3.2,
                "Trauma to perineum and vulva" : 2.1,
                "Normal pregnancy and/or delivery" : 2,
                "Umbilical cord complication" : 2.1,
                "Forceps delivery" : 2.2,
                "Administrative/social admission" : 4.2,
                "Prolonged pregnancy" : 2.4,
                "Other complications of pregnancy" : 2.5
                }

    los_a = []
    for dx, los in avg_los.items():
        for first in dx:
            los_a.append([dx, los])   
    columns=['dx','los']
    los_df= pd.DataFrame(los_a,columns=columns)
    patients = [['Boal', 'Medical examination/evaluation', 1.1],
        ['Boal', 'Other complications of pregnancy', 3.3],
        ['Jones', 'Liveborn', 3.2],
        ['Ashbury', 'Forceps delivery', 2.0]]
    columns2=["pt", 'dx', "los_actual"]
    pt_df=pd.DataFrame(patients, columns=columns2)
    pt_los_df= pd.merge(pt_df, los_df,on='dx',how='inner' )
    pt_los_df['dif']=pt_los_df['los']-pt_los_df['los_actual']
    print(pt_los_df['dif'])
    name= pt_los_df['pt'].tolist()
    los= pt_los_df['los'].tolist()
    diff=pt_los_df['dif'].tolist()
    patient_los=[]
    patient_los.append(los)
    patient_los.append(name)
    patient_los.append(diff)
    print(patient_los)
    for diff in patient_los:
        if diff==0:
            conclusion="just right"
        elif diff>0:
                conclusions= 'too long'
        else:
                conclusion="too short"
    patient_los.append(conlusion)
    ### END SOLUTION
    print(conclusion) 
    #I could not get the code for th conclusions to work but wanted to turn in something
```

```python

```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

---

```python
%%bash
git pull
git add week04_assignment_2.ipynb
git commit -a -m "Submitting the week 4 programming assignment"
git push
```

Output:
```
Already up to date.
[main dbbfdb1] Submitting the week 4 programming assignment
 2 files changed, 631 insertions(+), 4 deletions(-)
 create mode 100644 week04/week04_assignment_2.ipynb
To github.com:skuca/hds5210-2021.git
   de317d6..dbbfdb1  main -> main
```


---

If the message above says something like _Submitting the week 4 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.