# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

```python
### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.
```

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['one', 'two', 'three', 'four'])
    2
    """
    
    longest_pos = 0
    longest_len = 0
    
    print("Position  Item         Longest So Far")
    print("--------- ------------ --------------")
    for position, item in enumerate(strings):
        
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))
            
    return longest_pos

### END SOLUTION
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

```python
def longest(L):
    """(list) -> int
    Finds the longest string in this list.
    
    >>> longest(['my','name','is','Paul Boal'])
    3
    
    >>> longest("happy birthday to me".split(" "))
    1
    """
    
    position = 0
    length = 0
    
    for pos, word in enumerate(L):
        if len(word) > length:
            position = pos
            length = len(word)
    
    return position
    
    
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

### 25.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION

def f_to_c(temp_f):
    """(float) -> float
    Converts Farenheit to Celsius
    
    >>> f_to_c(212)
    100.0
    
    >>> f_to_c(32)
    0.0
    
    >>> f_to_c(98.6)
    37.0
    """
    
    return (5/9 * (temp_f - 32))

### END SOLUTION
```

```python
### BEGIN SOLUTION
def f_to_c(F):
    """(float) -> float
    
    >>> f_to_c(212)
    100.0
        
    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0
    """
    return (5/9) * (F-32)

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
f_to_c(212)
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
def los(admit, discharge, diagnosis):
    """(date, date, str) -> int
    Return length of stay in days based on the rules provided.
    If diagnosis is "Observation" then LOS = 0
    Otherwise LOS is the number of whole days between admin and discharge dates.
    
    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5
    
    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0
    """
    d = (discharge - admit).days
    
    if diagnosis == 'Observation':
        return 0
    else:
        return d
    
### END SOLUTION
```

```python
def los(admit, discharge, diagnosis):
    """(date, str, date)->int
    Computes length of stay in days.  If diagnosis is "Observation" then set the length of stay to 0.

    >>> los(date(2021,1,3), date(2021,1,8), 'COPD')
    5
    
    >>> los(date(2021,1,5), date(2021,1,9), 'Hypertension')
    4
    
    >>> los(date(2021,1,23), date(2021,1,25), 'Observation')
    0
    """
    if diagnosis == "Observation":
        return 0
    else:
        return (discharge - admit).days
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

```python
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for e, v in encounters.items():
    print(los(v['admit'],v['discharge'],v['diagnosis']))
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
from datetime import date
def average_los(encounter):
    """{dict}-> float
    Calculate the los by subtracting the admit date from the discharge date. Then sum for total LOS. Count number of encounters excluding encounters
    with los less than 1. To calculate average lenght of stay, divide total los by total encounters.
    """
    encounters = { "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
        "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
        "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
        "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
        "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}}
   
    los =[]
    
    for val in encounters.values():
        print(val)
     #didn't print any values  
    for los_a in encounters.values(['admit', 'discharge']):
        los_a= encounters['discharge']- encounters['admit']
    los.append(los_a) 
    print(los)
    total_los=sum(los)
    total_enc= len(los)
    average_los=total_enc/total_los
    return average_los({})
## not sure why the cod is not working- I did try- Attmept number 2 is below.
 

```

```python
   
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-47-008d658f0ab9> in <module>
----> 1 total_los=sum(los)
      2 total_enc= len(los)
      3 average_los=total_enc/total_los
      4 return average_los({})
      5 

TypeError: 'function' object is not iterable
```

```python
#the code worked, but did not return the proper format
### BEGIN SOLUTION
from datetime import date
import pandas as pd
import numpy as np
def average_los(encounter):
    """({dict})-> float
    Calculate the los by subtracting the admit date from the discharge date. Then sum for total LOS. Count number of encounters excluding encounters
    with los less than 1. To calculate average lenght of stay, divide total los by total encounters.
    
    >>> {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8)}
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,15) }}
  
    5
    
    >>> {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8)}
          {"E1236": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,7) }}  
    4.5
    
    """
    encounters = { "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
            "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
            "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
            "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
            "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}}

    enc_df= pd.DataFrame.from_dict(encounters) 
    enc_df = enc_df.transpose()
    enc_df['los'] = enc_df['discharge']-enc_df['admit']
    enc_df['los']= enc_df['los'].dt.days.astype(float)
    encounters= enc_df.to_dict()
    los_r = encounters['los'].values()
    total_los=sum(los_r)
    c_enc=len([elm for elm in los_r if elm >0])
    average_los=total_los/c_enc
    return average_los({})
    print(average_los)

    
 


   

### END SOLUTION
```

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}}


assert(average_los(encounters)) == 2.75
```

Output:
```
<ipython-input-56-cabab2b8bbc4> in <module>
      8 
      9 
---> 10 assert(average_los(encounters)) == 2.75

<ipython-input-55-0508fcdf5d57> in average_los(encounter)
     34     c_enc=len([elm for elm in los_r if elm >0])
     35     average_los=total_los/c_enc
---> 36     return average_los({})
     37     print(average_los)
     38 

TypeError: 'float' object is not callable
```

```python

```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8)}
Expecting:
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,15) }}
**********************************************************************
File "__main__", line 11, in NoName
Failed example:
    {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8)}
Exception raised:
    Traceback (most recent call last):
      File "/opt/tljh/user/lib/python3.6/doctest.py", line 1330, in __run
        compileflags, 1), test.globs)
      File "<doctest NoName[0]>", line 1
        {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8)}
                                                                                             ^
    SyntaxError: unexpected EOF while parsing
Trying:
    {"E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8)}
Expect
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python

### BEGIN SOLUTION
def c_to_f(temp_c):
# F=(C x 9/5) + 32
    """(float)->float
    Converts temperture in celcius to farenheit
    >>> c_to_f(100.0)
    212.0
    
    >>> c_to_f (0.0)
    32.0
    
    >>> round(c_to_f(37),2)
    98.6
    """
    return (temp_c*(9/5))+ 32
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

---

```python
%%bash
git pull
git add week05_assignment_2.ipynb
git commit -a -m "Submitting the week 5 programming assignment"
git push
```


---

If the message above says something like _Submitting the week 5 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.

```python

```